# LAB 0: Test Submission [Total: 3 points]

### Important

Please ensure that you run the cell below before running any others. This will download all required files and packages, as well as installing the necessary packages to ensure the code runs successfully. If you restart the kernel or your runtime session (in Colab), be sure to rerun this cell before running any others.

```python
required_files = "https://github.com/mainuddin-rony/inst447-fall2024/raw/main/assignment/lab/lab0/required_files.zip"
! wget $required_files && unzip -j required_files.zip
! mkdir tests && mv q1.py tests
! pip install -r requirements.txt
```

```python
# Initialize Otter
import otter
grader = otter.Notebook()
```

The purpose of this assignment is to familiarize with the submission process of this course. This assignment comprises a simple question for **manual grading** and a basic program for **Autograder** evaluation.

## Manually graded [1 point]

Write down your name in the following cell.

```python
## Your name here
name = ""
name
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

## Autograded Task [2 points]

### Q1: Calculate mean and median of a numpy array

Define a function called `array_cal` that takes two parameters - a list of integers and an operation name. The function should first convert the list to a numpy array, perform the operation on the array, and return the result.

For example, given the list `[1,2,3]` and the operation `mean`, it should return `2.0`.

**Parameters**
* `num_lst` &ndash; a list of integers,
* `op` &ndash; name of the operation string (either `mean` or `median`)

**Returns**
* A numpy float value which is the result after performing the operation

**Hint**
You can use `numpy.mean()` and `numpy.median()` function to calculate mean and median

---

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

```python
import numpy as np
def array_cal(num_lst, op):
...

# Test with different arguments!
num_lst = [1,2,3]
op = 'mean'
result = array_cal(num_lst, op)
print("The " + op + " array is " + str(result))
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q1")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

## Submission

Don't forget to run all cells in your notebook and then save it. To save, click on `File`, then select `Save/Save Notebook`. After that, download the notebook by going to `File`, then `Download` (for Anaconda notebook), and choosing `Download .ipynb` (for Colab). Finally, submit the notebook on `Gradescope` using the link found on ELMS.

```python

```

Tests for q1 (points: None)
```
>>> assert array_cal([1, 2, 3], 'mean') == 2.0, "For input [1,2,3] and 'mean'', the function should return 2.0"
```
```
>>> assert array_cal([5, 6, 7], 'mean') == 6.0, "For input [5,6,7] and 'mean'', the function should return 3.0"
```
```
>>> assert array_cal([5, 6, 7], 'median') == 6.0, "For input [5,6,7] and 'median'', the function should return 6.0"
```
```
>>> assert array_cal([5, 6, 4, 5, 7, 3, 7], 'median') == 5.0, "For input [5,6,4,5,7,3,7] and 'median'', the function should return 5.0"
```
